rule_files:
- rules.yaml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-api"}
    values: "0x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-api"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPIDegraded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        deployment: mistral-api
        exported_namespace: default
      exp_annotations:
        description: Mistral API degraded, there are unavailable API pods
        summary: Mistral API degraded

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-api"}
    values: "1x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-api"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPIDegraded
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-api"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPIDown
    exp_alerts:
    - exp_labels:
        severity: critical
        namespace: default
        service: mistral-montemplate
        deployment: mistral-api
        exported_namespace: default
      exp_annotations:
        description: Mistral API down, there are no available API pods
        summary: Mistral API down

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-api"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPIDown
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-api"}
    values: "60+60x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-api", resource="cpu"}
    values: "1x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-api", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPICPULoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral API use 90 percent of CPU limit
        summary: Mistral API use 90 percent of CPU limit

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-api"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-api", resource="cpu"}
    values: "1x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-api", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPICPULoad
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-api"}
    values: "1x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-api", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPIMemoryLoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral API use 90 percent of memory limit
        summary: Mistral API use 90 percent of memory limit

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-api"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-api", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralAPIMemoryLoad
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-engine"}
    values: "0x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-engine"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineDegraded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        exported_namespace: default
        deployment: mistral-engine
      exp_annotations:
        description: Mistral Engine degraded, there are unavailable Engine pods
        summary: Mistral Engine degraded

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-engine"}
    values: "1x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-engine"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineDegraded
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-engine"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineDown
    exp_alerts:
    - exp_labels:
        severity: critical
        namespace: default
        service: mistral-montemplate
        exported_namespace: default
        deployment: mistral-engine
      exp_annotations:
        description: Mistral Engine down, there are no available Engine pods
        summary: Mistral Engine down

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-engine"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineDown
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: rabbitmq_queues{namespace="default",queue="_mistral_engine"}
    values: "51x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineOverloaded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        queue: _mistral_engine
      exp_annotations:
        description: Mistral Engine overloaded, there are lots of unprocessed RPC messages
        summary: Mistral Engine overloaded

- interval: 1m
  input_series:
  - series: rabbitmq_queues{namespace="default",queue="_mistral_engine"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineOverloaded
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-engine"}
    values: "60+60x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-engine", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineCPULoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Engine use 90 percent of CPU limit
        summary: Mistral Engine use 90 percent of CPU limit

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-engine"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-engine", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineCPULoad
    exp_alerts: []

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-engine"}
    values: "1x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-engine", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineMemoryLoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Engine use 90 percent of memory limit
        summary:  Mistral Engine use 90 percent of memory limit

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-engine"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-engine", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralEngineMemoryLoad
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-executor"}
    values: "0x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-executor"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorDegraded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        exported_namespace: default
        deployment: mistral-executor
      exp_annotations:
        description: Mistral Executor degraded, there are unavailable Executor pods
        summary:  Mistral Executor degraded

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-executor"}
    values: "1x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-executor"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorDegraded
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-executor"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorDown
    exp_alerts:
    - exp_labels:
        severity: critical
        namespace: default
        service: mistral-montemplate
        exported_namespace: default
        deployment: mistral-executor
      exp_annotations:
        description: Mistral Executor down, there are no available Executor pods
        summary:  Mistral Executor down

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-executor"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorDown
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: rabbitmq_queues{namespace="default",queue="_mistral_executor"}
    values: "51x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorOverloaded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        queue: _mistral_executor
      exp_annotations:
        description: Mistral Executor overloaded, there are lots of unprocessed RPC messages
        summary:  Mistral Executor overloaded

- interval: 1m
  input_series:
  - series: rabbitmq_queues{namespace="default",queue="_mistral_executor"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorOverloaded
    exp_alerts: []  

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-executor"}
    values: "60+60x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-executor", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorCPULoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Executor use 90 percent of CPU limit
        summary: Mistral Executor use 90 percent of CPU limit

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-executor"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-executor", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorCPULoad
    exp_alerts: []

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-executor"}
    values: "1x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-executor", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorMemoryLoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Executor use 90 percent of memory limit
        summary: Mistral Executor use 90 percent of memory limit

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-executor"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-executor", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralExecutorMemoryLoad
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-notifier"}
    values: "0x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-notifier"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierDegraded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        exported_namespace: default
        deployment: mistral-notifier
      exp_annotations:
        description: Mistral Notifier degraded, there are unavailable Notifier pods
        summary: Mistral Notifier degraded

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-notifier"}
    values: "1x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-notifier"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierDegraded
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-notifier"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierDown
    exp_alerts:
    - exp_labels:
        severity: critical
        namespace: default
        service: mistral-montemplate
        exported_namespace: default
        deployment: mistral-notifier
      exp_annotations:
        description: Mistral Notifier down, there are no available Notifier pods
        summary: Mistral Notifier down

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-notifier"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierDown
    exp_alerts: []

- interval: 1m
  input_series:
  - series: rabbitmq_queues{namespace="default",queue="_mistral_notifier"}
    values: "51x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierOverloaded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        queue: _mistral_notifier
      exp_annotations:
        description: Mistral Notifier overloaded, there are lots of unprocessed RPC messages
        summary: Mistral Notifier overloaded

- interval: 1m
  input_series:
  - series: rabbitmq_queues{namespace="default",queue="_mistral_notifier"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierOverloaded
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-notifier"}
    values: "60+60x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-notifier", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierCPULoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Notifier use 90 percent of CPU limit
        summary: Mistral Notifier use 90 percent of CPU limit

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-notifier"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-notifier", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierCPULoad
    exp_alerts: []

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-notifier"}
    values: "1x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-notifier", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierMemoryLoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Notifier use 90 percent of memory limit
        summary: Mistral Notifier use 90 percent of memory limit

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-notifier"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-notifier", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralNotifierMemoryLoad
    exp_alerts: []

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-monitoring"}
    values: "0x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-monitoring"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringDegraded
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
        exported_namespace: default 
        deployment: mistral-monitoring
      exp_annotations:
        description: Mistral Monitoring degraded, there are unavailable Monitoring pods
        summary: Mistral Monitoring degraded

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-monitoring"}
    values: "1x5"
  - series: kube_deployment_status_replicas{exported_namespace="default", deployment="mistral-monitoring"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringDegraded
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-monitoring"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringDown
    exp_alerts:
    - exp_labels:
        severity: critical
        namespace: default
        service: mistral-montemplate
        exported_namespace: default 
        deployment: mistral-monitoring
      exp_annotations:
        description: Mistral Monitoring down, there are no available Monitoring pods
        summary: Mistral Monitoring down

- interval: 1m
  input_series:
  - series: kube_deployment_status_replicas_available{exported_namespace="default", deployment="mistral-monitoring"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringDown
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-monitoring"}
    values: "60+60x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-monitoring", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringCPULoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Monitoring use 90 percent of CPU limit
        summary: Mistral Monitoring use 90 percent of CPU limit

- interval: 1m
  input_series:
  - series: container_cpu_usage_seconds_total{namespace="default", container="mistral-monitoring"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-monitoring", resource="cpu"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringCPULoad
    exp_alerts: []
    
- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-monitoring"}
    values: "1x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-monitoring", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringMemoryLoad
    exp_alerts:
    - exp_labels:
        severity: warning
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: Mistral Monitoring use 90 percent of memory limit
        summary: Mistral Monitoring use 90 percent of memory limit

- interval: 1m
  input_series:
  - series: container_memory_usage_bytes{namespace="default", container="mistral-monitoring"}
    values: "0x5"
  - series: kube_pod_container_resource_limits{namespace="default", container="mistral-monitoring", resource="memory"}
    values: "1x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: MistralMonitoringMemoryLoad
    exp_alerts: []

- interval: 1m
  input_series:
  - series: rabbitmq_status{namespace="default"}
    values: "2x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: RabbitMQDown
    exp_alerts:
    - exp_labels:
        severity: critical
        namespace: default
        service: mistral-montemplate
      exp_annotations:
        description: RabbitMQ down
        summary: RabbitMQ down

- interval: 1m
  input_series:
  - series: rabbitmq_status{namespace="default"}
    values: "0x5"
  alert_rule_test:
  - eval_time: 5m
    groupname: default-mistral-montemplate
    alertname: RabbitMQDown
    exp_alerts: []
